# Assignment 3: Book Recommenders II

This notebook provides starter code and a template for formatting your work. While some basic instructions are included, please refer to the full instructions on MyCourses for details and guidance on how to approach each task. 

Note that each response cell is followed by one or more test cells. You can use these to help test your code. Note that failing a test does not necessarily mean that your code is wrong. Notably, some tasks ask you to print output to the screen. The test cell will check that your solution exactly matches mine. Solutions that are correct but have small formatting differences will fail, but would be awarded full marks during grading. It is also possible for your code to pass the test cells, even if there are some errors that will be detected during grading. **However, despite these limitations the test cell can provide important information to help you check the correctness of your solution**

Please write your name in the cell below.

<!-- grade_id: name, points: 0, grade: true, solution: true -->
Andy Salazar Parra

### Load libraries

You must run the following cell for the autograder test cells to work. If you restart the notebook you will need to run this cell again: 

```python
try: from nose.tools import assert_equal
except ImportError as e:
    %pip install nose
    from nose.tools import assert_equal

try: from unittest import mock
except ImportError as e:
    %pip install unittest
    from unittest import mock
    from unittest.mock import patch

try: import io
except ImportError as e:
    %pip install io
    import io

try: import sys
except ImportError as e:
    %pip install sys
    import sys
    
try: import os
except ImportError as e:
    %pip install os
    import os
```


***
# Part 1: File Management

The first part of this assignment will be to create functions that enable the customer and book databases to read and written to file. I have provided the files `books.txt` and `ratings.txt`, which will enable you to work with a larger number of books and customer ratings. This is needed for the sophisitacted recommendation algorithm to produce reasonable results. You will need to download these files from MyCourses and store them in the same directory as this notebook file. 

We will write a total of three functions in this part. The first two will read in the data files and initialize data structures for handling our book and customer "databases". The third will write out the ratings data back to the file at the end of the session so that new ratings are added. *Note you should keep a backup of the original ratings file close at hand!* You do not need to write out the book file as we will be working with a fixed set of books. *(Extending our code to handle the ability to add books would be a nice extension if you are looking for a challenge, but it's non-trivial.)* 


## Task 1: Write `readBooks()`
First, you will write a function `readBooks()` that reads the file books.txt and initialize the books database as a dictionary.

Recall from the last assignment that we'll use a dictionary, using the ISBN number as the key, and making the value a tuple in the form `(YEAR, TITLE, AUTHOR, INDEX)`. The ISBN, publication year, book title, and author are all in `books.txt` with one line per book. Index should simply be the line number, counting from 0. That is,

    0 The Hitchhiker's Guide To The Galaxy
    1 Watership Down
    ...
    54 The Chrysalids
    
As the ratings scores are in the same order, keeping an index will make it a little easier to look up a particular book's rating in the customer data. *Note that there are other ways that this dictionary could be structured that would lead to an equally viable final solution. However, it is often the case that multiple programmers will work together on a project and will need to agree on how they will organize their data structures. As such, it is not entirely unrealistic that I am imposing a particular design.*

### Pseudocode

* Initialize a dictionary to hold our books database
* Open the file specified by the input parameter `filename`
* Loop over each line in the file
    * Split each line into each the individual components
    * Add each component (plus a line index) to the dictionary as described above
* return the books database

*Many students ignore the input parameter and just hardcode the filename into their function definition. This works because we only ever work with one dataset but for full marks your solution to this task must use a filename provided as an input parameter*

```python
def readBooks(filename):
    # YOUR CODE HERE
    # made book DB global so I could print it in the text function
    global bookDB
    # intialize book DB
    bookDB = {}
    # open file for reading
    file1 = open(filename, "r")
    # read lines and store them in list
    lines = file1.readlines()
    # iterate over file lines
    for i in lines:
        # separate lines into components
        component = i.split(", ")
        # remove \n at the end of last component
        component[3] = component[3].replace("\n", "")
        # get book index
        bookIndex = lines.index(i)
        # add isbn as key in dictionary and add the other comoponents as dict values
        bookDB[component[0]] = (component[1], component[2], component[3], bookIndex)
    # close the file
    file1.close()
    # return book dictionary
    return bookDB
```

<!-- grade_id: task1-autograder1, points: 2, grade: true, solution: false -->
```python
# Test your code -- Run this cell to check readBooks works with books.txt
sys._jupyter_stdout = sys.stdout
sys.stdout = open(os.devnull, 'w')    
try: assert_equal(readBooks("books.txt"), {'0517693119': ('1989', "The Hitchhiker's Guide To The Galaxy", 'Douglas Adams', 0), '0025002600': ('1978', 'Watership Down', 'Richard Adams', 1), '0937247065': ('1988', 'The Five People You Meet in Heaven', 'Mitch Albom', 2), '8809744454': ('2009', 'Speak', 'Laurie Halse Anderson', 3), '0929631048': ('1989', 'I Know Why the Caged Bird Sings', 'Maya Angelou', 4), '1595141712': ('2010', 'Thirteen Reasons Why', 'Jay Asher', 5), '0553382570': ('1964', 'Foundation Series', 'Isaac Asimov', 6), '0375890291': ('2002', 'The Sisterhood of the Travelling Pants', 'Ann Brashares', 7), '0731814909': ('2010', 'A Great and Terrible Beauty', 'Libba Bray', 8), '8495618818': ('2004', 'The Da Vinci Code', 'Dan Brown', 9), '0230037518': ('2008', 'The Princess Diaries', 'Meg Cabot', 10), '3641082505': ('2012', "Ender's Game", 'Orson Scott Card', 11), '0425116840': ('1989', 'The Hunt for Red October', 'Tom Clancy', 12), '0545229936': ('2009', 'The Hunger Games', 'Suzanne Collins', 13), '0743273565': ('1925', 'The Great Gatsby', 'F. Scott Fitzgerald', 14), '0142411736': ('2004', "Ranger's Apprentice Series", 'John Flanagan', 15), '0606338039': ('2005', 'Inkheart', 'Cornelia Funke', 16), '0441007465': ('2000', 'Neuromancer', 'William Gibson', 17), '0399501487': ('1963', 'Lord of the Flies', 'William Golding', 18), '0345439740': ('1957', 'The Princess Bride', 'William Goldman', 19), '1863330550': ('1992', 'Dinotopia: A Land Apart from Time', 'James Gurney', 20), '1894448200': ('2000', 'Far North', 'Will Hobbs', 21), '0425190374': ('2003', 'Practical Magic', 'Alice Hoffman', 22), '0060850523': ('1968', 'Brave New World', 'Aldous Huxley', 23), '0143176763': ('2010', 'The Summer Tree', 'Guy Gavriel Kay', 24), '0871295377': ('1969', 'Flowers For Algernon', 'Daniel Keyes', 25), '0618439102': ('2004', 'Owl in Love', 'Patrice Kindl', 26), '1569319006': ('2003', 'Naruto', 'Masashi Kishimoto', 27), '1591164418': ('2004', 'Bleach (graphic novel)', 'Tite Kubo', 28), '0230738034': ('2008', 'Kiss the Dust', 'Elizabeth Laird', 29), '0871299208': ('1970', 'To Kill a Mockingbird', 'Harper Lee', 30), '0064404990': ('1950', 'The Lion the Witch and the Wardrobe', 'C S Lewis', 31), '0312331754': ('2004', 'The Bourne Series', 'Robert Ludlum', 32), '0547416113': ('2002', 'Life of Pi', 'Yann Martel', 33), '0440240167': ('2010', 'Breathless', 'Lurlene McDaniel', 34), '1905654391': ('2008', 'Twilight Series', 'Stephenie Meyer', 35), '0061975133': ('2009', 'Sabriel', 'Garth Nix', 36), '1411469380': ('1982', 'Nineteen Eighty-Four (1984)', 'George Orwell', 37), '0375890369': ('2003', 'Eragon', 'Christopher Paolini', 38), '0330338730': ('1994', 'Hatchet', 'Gary Paulsen', 39), '1416549178': ('2007', "My Sister's Keeper", 'Jodi Picoult', 40), '0375846724': ('1997', 'The Golden Compass', 'Philip Pullman', 41), '0590353403': ('1998', 'Harry Potter Series', 'J.K. Rowling', 42), '8434878607': ('2001', 'Holes', 'Louis Sachar', 43), '1156393361': ('1968', 'Shonen Jump Series', 'Shueisha', 44), '0142300942': ('1988', 'The Shadow Club', 'Neil Shusterman', 45), '1443102067': ('1993', 'Bone Series', 'Jeff Smith', 46), '0394747232': ('1986', "Maus: A Survivor's Tale", 'Art Spiegelman', 47), '0143038095': ('1992', 'The Joy Luck Club', 'Amy Tan', 48), '0261102958': ('1992', 'The Lord of the Rings', 'J R R Tolkien', 49), '0345339683': ('1982', 'The Hobbit', 'J R R Tolkien', 50), '0143181538': ('2007', 'Shattered', 'Eric Walters', 51), '1603034013': ('1898', 'The War Of The Worlds', 'H G Wells', 52), '0547538648': ('2015', 'Dealing with Dragons', 'Patricia C. Wrede', 53), '0141038462': ('1964', 'The Chrysalids', 'John Wyndham', 54)})
finally: sys.stdout = sys._jupyter_stdout 
print("Test passed")
```

Output:
```
Test passed
```

<!-- grade_id: task1-autograder2, points: 1, grade: true, solution: false -->
```python
# Test your code -- Run this cell to check that it works with a file other than books.txt
sys._jupyter_stdout = sys.stdout
sys.stdout = open(os.devnull, 'w')    
with open('mtf.txt', 'w') as fileObj: fileObj.write("01, 1900, abc, def\n02, 1901, ghi, klm")
try: assert_equal(readBooks("mtf.txt"), {'01': ('1900', 'abc', 'def', 0), '02': ('1901', 'ghi', 'klm', 1)})
finally: sys.stdout = sys._jupyter_stdout 
if os.path.exists("mtf.txt"): os.remove("mtf.txt")
print("Test passed")
```

Output:
```
Test passed
```

### Test `readBooks()`

If you run the cell above with your implementation of `readBooks()` nothing should happen because the cell above should only contain a definition for `readBooks()`, it shouldn't call it. Just like in the last assignment, we will separate the tasks of defining our functions and testing them. This will enable running the cells with the function definitions, without calling the function, which will help both you in completing the assignment and me in grading it. 

Write a small function that tests your code by calling `readBooks()` and then printing out each key, value pair in the returned dictionary. Check that this is consistent with `books.txt` and your expectations. Once you are satisfied it does, you can comment out the line calling `testReadBooks()` so that this code doesn't run every time you run the entire file. As this code is just to test that `readBooks()` is working properly, you do not need to worry about making the output pretty (i.e., you can just let Python handle printing the tuple). The first four lines of your output should look like this: 

    0517693119 --> ('1989', "The Hitchhiker's Guide To The Galaxy", 'Douglas Adams', 0)
    0025002600 --> ('1978', 'Watership Down', 'Richard Adams', 1)
    0937247065 --> ('1988', 'The Five People You Meet in Heaven', 'Mitch Albom', 2)
    8809744454 --> ('2009', 'Speak', 'Laurie Halse Anderson', 3)
    
And the last two should look like this: 

    0547538648 --> ('2015', 'Dealing with Dragons', 'Patricia C. Wrede', 53)
    0141038462 --> ('1964', 'The Chrysalids', 'John Wyndham', 54)

```python
def testReadBooks(): 
    # YOUR CODE HERE
    readBooks("books.txt")
    for i in bookDB.items():
        print(i)

# testReadBooks()     # Comment out this line when you are done testing 
```

Output:
```
('0517693119', ('1989', "The Hitchhiker's Guide To The Galaxy", 'Douglas Adams', 0))
('0025002600', ('1978', 'Watership Down', 'Richard Adams', 1))
('0937247065', ('1988', 'The Five People You Meet in Heaven', 'Mitch Albom', 2))
('8809744454', ('2009', 'Speak', 'Laurie Halse Anderson', 3))
('0929631048', ('1989', 'I Know Why the Caged Bird Sings', 'Maya Angelou', 4))
('1595141712', ('2010', 'Thirteen Reasons Why', 'Jay Asher', 5))
('0553382570', ('1964', 'Foundation Series', 'Isaac Asimov', 6))
('0375890291', ('2002', 'The Sisterhood of the Travelling Pants', 'Ann Brashares', 7))
('0731814909', ('2010', 'A Great and Terrible Beauty', 'Libba Bray', 8))
('8495618818', ('2004', 'The Da Vinci Code', 'Dan Brown', 9))
('0230037518', ('2008', 'The Princess Diaries', 'Meg Cabot', 10))
('3641082505', ('2012', "Ender's Game", 'Orson Scott Card', 11))
('0425116840', ('1989', 'The Hunt for Red October', 'Tom Clancy', 12))
('0545229936', ('2009', 'The Hunger Games', 'Suzanne Collins', 13))
('074327
```

<!-- grade_id: task1t-autograder, points: 1, grade: true, solution: false -->
```python
# Test your code
sys._jupyter_stdout = sys.stdout
sys.stdout = open(os.devnull, 'w')    

old = readBooks
del readBooks

try: testReadBooks()
except NameError: pass
else: 
    raise AssertionError("testReadBooks() does not call readBooks()")
finally:
    sys.stdout = sys._jupyter_stdout
    readBooks = old
    del old

print("Test passed")
```

Output:
```
Test passed
```

### Write `printTen()`

Write a second test function `printTen()`. This function should likewise call `readBooks()` to initialize a dictionary database. It should then print 10 books titles at random. Your resulting output might look something like this (though the exact books will be selected at random):

    The Hunger Games by Suzanne Collins
    Far North by Will Hobbs
    Sabriel by Garth Nix
    The Bourne Series by Robert Ludlum
    Holes by Louis Sachar
    The Summer Tree by Guy Gavriel Kay
    Neuromancer by William Gibson
    Flowers For Algernon by Daniel Keyes
    The Sisterhood of the Travelling Pants by Ann Brashares
    Kiss the Dust by Elizabeth Laird

*Hint: You will need to use the random module which I've already imported in the skeleton code below. random has a method sample which should come in handy. random.sample(range(0,100), 10) returns 10 random numbers between 0 and 99 inclusive.*

<!-- grade_id: task1b-answer, points: 0, grade: false, solution: true -->
```python
import random
def printTen():
    # YOUR CODE HERE
    # read file
    readBooks("books.txt")
    # get random list of ten books
    tenBooks = (random.sample(bookDB.items(),10))
    # iterate over list and print book info for each item
    for i in tenBooks:
        print((i[1][1]), "("+(i[1][0])+")", "by", (i[1][2]))

#printTen() #Comment out this line when you are done testing 
```

Output:
```
The Sisterhood of the Travelling Pants (2002) by Ann Brashares
Inkheart (2005) by Cornelia Funke
Flowers For Algernon (1969) by Daniel Keyes
The Summer Tree (2010) by Guy Gavriel Kay
Practical Magic (2003) by Alice Hoffman
The Hunt for Red October (1989) by Tom Clancy
Ender's Game (2012) by Orson Scott Card
Shonen Jump Series (1968) by Shueisha
Twilight Series (2008) by Stephenie Meyer
Neuromancer (2000) by William Gibson
```

<!-- grade_id: task1b-autograder1, points: 1, grade: true, solution: false -->
```python
# Test your code 1
sys._jupyter_stdout = sys.stdout
sys.stdout = open(os.devnull, 'w')    

old = random.sample
del random.sample

try: printTen()
except AttributeError: pass
else: 
    raise AssertionError("printTen() does not call random.sample()")
finally:
    sys.stdout = sys._jupyter_stdout
    random.sample = old
    del old

print("Test passed")
```

Output:
```
Test passed
```

<!-- grade_id: task1b-autograder2, points: 1, grade: true, solution: false -->
```python
# Test your code 2
sys._jupyter_stdout = sys.stdout
sys.stdout = open(os.devnull, 'w')    

old = readBooks
del readBooks

try: printTen()
except NameError: pass
else: 
    raise AssertionError("printTen() does not call readBooks()")
finally:
    sys.stdout = sys._jupyter_stdout
    readBooks = old
    del old

print("Test passed")
```

Output:
```
Test passed
```

## Task 2: Write `readRatings()`

Now, you will write a function `readRatings()` that reads the file ratings.txt and initializes the customer database. As for books, we will also use a dictionary to store our customer data, but it will be a bit simpler. 

The customer's ID will be the dictionary key. All the customer names in `ratings.txt` are unique so we know that when we initially load the database from file all the customers will have a unique ID. Moreover, when we wrote `openAccount()` in the last assignment we also made sure to check that the user ID was unique before accepting it (and adding the person to the database), which will ensure customer IDs remain unique as new customers are added to the database. 

The dictionary value will be a list of book ratings. Note that when we read in the data from file, we will initially have this data as a string, you will need to convert the string to a list and then each item in the list to an integer. This is a bit of work now, but will make many later tasks easier (since we will often want to compute things from our ratings&mdash;as we already saw in the previous assignment when writing `getAverageRatings()`)

### Pseudocode

* Initialize a dictionary to hold our customers database
* Open the file specified by the input parameter `filename` for reading
* Loop over each line in the file
    * Split each line into customer id and a string of ratings
    * Split the string of ratings into a list of strings
    * Convert the list of strings to a list of numbers 
        * Break the string into a list of string ratings
        * Loop over each element of this list and replace its string value with an integer equivalent
    * Add the customer id and the list of ratings to the dictionary as a key, value pair
* return the customer database


<!-- grade_id: task2-answer, points: 0, grade: false, solution: true -->
```python
def readRatings(filename):
    # made customerDB global so it could be called in test function
    global customerDB
    # intialize customer dictionary
    customerDB = {}
    # YOUR CODE HERE
    # open file for reading
    file2 = open(filename, "r")
    # read file lines and store them in list
    lines = file2.readlines()
    # iterate over list of lines
    for i in lines:
        # separate line into compunents
        component = i.split(", ")
        # create username value from component 1
        username = component[0]
        # create list of ratings from component 2
        ratings = (component[1]).split()
        # iterate over ratings list
        for i in range (0, len(ratings)):
            # change ratings to integer
            ratings[i] = int(ratings[i])
        # add ratings as value and username as key in customerDB
        customerDB[username] = (ratings)
    # close file
    file2.close()
    # return customer DB
    return customerDB
    
```

<!-- grade_id: task2-autograder1, points: 2, grade: true, solution: false -->
```python
# Test your code -- check readRatings works with ratings.txt
sys._jupyter_stdout = sys.stdout
sys.stdout = open(os.devnull, 'w')    
try: assert_equal(readRatings("ratings.txt"), {'Ben': [5, 0, 0, 0, 0, 0, 0, 1, 0, 1, -3, 5, 0, 0, 0, 5, 5, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 1, 3, 0, 1, 0, -5, 0, 0, 5, 5, 0, 5, 5, 5, 0, 5, 5, 0, 0, 0, 5, 5, 5, 5, -5], 'Moose': [5, 5, 0, 0, 0, 0, 3, 0, 0, 1, 0, 5, 3, 0, 5, 0, 3, 3, 5, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 3, 5, 0, 0, 0, 0, 0, 5, -3, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 5, 5, 0, 3, 0, 0], 'Reuven': [5, -5, 0, 0, 0, 0, -3, -5, 0, 1, -5, 5, 0, 1, 0, 1, -3, 1, -5, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, -5, 1, 0, 1, 0, -5, 0, 3, -3, 3, 0, 1, 5, 1, 0, 0, 0, 0, 0, 1, 3, 1, 5, 1, 3], 'Cust1': [3, 3, 5, 0, 0, 0, 3, 0, 0, 3, 0, 3, 0, 0, 0, 0, 0, 3, 0, 5, 0, 0, 0, 1, 3, 1, 0, 0, 0, 0, 0, 3, 0, 3, 0, 0, 0, 1, 3, 0, 0, 3, 3, 0, 0, 0, 5, 0, 0, 3, 1, 0, 0, 0, 0], 'Cust2': [3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 3, 1, 0, 0, 0, 3, 0, 0, 0, 3, 0, 3, 3, 5, 0, 3, 0, 3], 'Francois': [3, 3, 5, 0, 0, 0, 3, 0, 0, 3, 0, 3, 0, 0, 0, 0, 0, 3, 0, 5, 0, 0, 0, 1, 3, 1, 0, 0, 0, 0, 0, 3, 0, 3, 0, 0, 0, 1, 3, 0, 0, 3, 3, 0, 0, 0, 5, 0, 0, 3, 1, 0, 0, 0, 0], 'Jim C': [5, 3, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 3, 0, 0, 0, 0, 0, 5, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 3, 3, 0, 0, 0, 0], 'Iren': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 1, 0, 0, 0, 1, 0, 0, 5, 0, 0, 0, 5, 3, 5, 0, 0, 0, 0, 5, 3, 0, 3, 0, -5], 'Cust3': [0, 0, 0, -3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 1, 0, -3, 5, 3, 0, 0, 1, 0, 0, 0, 1, 0, 0, 0, 0], 'Cust4': [0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, -5, 0, 0, -5, 5, 0, 0, 0, 0, 0, 0, 0, -3, 0, 0, 0, -3], 'Cust5': [0, -3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 5, 3, 0, 0, 0, 0, 0, 5, 0, 3, 0, 3, 5, 3, 0, 0, 0, 3, 0, 5, 5, 0, 0, 5, 0], 'Cust6': [1, 1, 3, 3, -3, 3, 3, 5, 5, 5, 5, 1, 3, 1, 1, 0, 5, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 3, 0, 3, 0, 3, 5, 5, 5, 0, -3, 0, 0, 0, 0, 0, 0, 3, 1, 1], 'Cust7': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, -3, 0, 0, 5, 5, 0, 0, 5, 0, 0, 3, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 5], 'Cust8': [3, 0, 0, 5, 0, 0, 0, 5, 0, 5, 3, 0, 0, 0, 0, 0, 3, 0, 5, 5, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 5, 5, 3, 0, 0, 5, 0, 0, 5, 5, 5, 3, 5, 5, 0, 0, 5, 0, 0, 5, 5, 0, 5, 0, 5], 'Cust9': [5, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, -3, 0, 0, 0, 0, 3], 'andrew': [3, 3, 1, 0, 3, 0, 3, 0, 0, -3, 0, 5, 3, 0, 1, 0, 0, 5, 3, 0, 0, 0, 0, 1, 0, 3, 0, 1, 0, 0, 3, 5, 3, 3, 0, 0, 0, 5, 0, 5, 0, 3, 3, 0, -3, 0, 0, 5, 1, 5, 3, 0, 3, 0, 0], 'ma': [0, 0, 0, 0, 5, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, -3, 0, 0, 0, 0, -3, 0, 0, 0, 0, 0, 0, 0, 3, 0, 3, 0, 0, 0, -3, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, -3, -3, 0, 0, 0, 0], 'matt c': [0, 3, 0, 0, 0, 0, 0, 0, 0, 5, 0, 5, 0, 0, -3, 5, 0, 0, 1, 0, 0, 3, 0, 0, 3, 0, 0, 0, 0, 0, 1, 3, 3, 0, 0, 0, 0, 0, 5, 0, 0, 0, 5, 3, 0, 3, 0, 0, 0, 5, 5, 3, 0, 0, 0], 'Ocelot': [5, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 3, 0, 0, 5, 3, 0, 0, 0, 0, 3, 0, 1, 1, 0, 0, 0, 3, 3, 0, 5, 0, -3, 0, 3, 0, 0, 0, 5, 3, 0, 0, 0, 0, 0, 0, 0, 3, 0, 1, 0, 0], 'Cust10': [5, 0, 0, 0, 1, 0, 0, 0, 0, -5, 0, 5, 0, 0, 0, 0, 0, 3, -3, 3, 0, 0, 0, 5, 0, 5, 0, 0, 0, 0, 3, 3, 0, 5, 0, 0, 0, 5, 0, 0, 0, 0, 3, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0], 'crom': [0, 0, 3, 0, 0, 0, 0, 0, 0, -3, 0, 0, 0, 0, 5, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 3, 0, 0, 0, 0], 'joe': [3, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 5, 0, 0, 5, 0, 1, 1, 3, 0, 1, 0, 3, 0, 0, 0, 0, 3, 3, 0, 5, 3, 1, 3, 0, 1, 0, 3, 1, 0, 0, 5, 3, 0, 5, 0, 1, 0, 0, 1, 5, 0, 0, 0, 0], 'Priscilla': [5, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 3, 3, 0, 1, 0, 0, 0, 0, 0, 0, 0, 3, 5, 0, 0, 0, 0, 0, 0, 0, 1, 0, 3, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, -3, 1, 0, 1, 0, -3], 'Harry Potter': [0, 0, 0, 0, 0, 0, 0, 0, 0, -5, 0, 1, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 1, 1, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0], 'I. Ned': [0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, -3, 0, 0, 0, 0, 0, 0, 1, 1, 0, 0, 0, -3, 0, 0, 3, 1, 0, 3, 1, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, -3], 'ender': [5, 0, 0, 0, 0, 0, 5, 0, 0, 3, 0, 5, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 3, 0, 5, 0, 3], 'Apollo': [0, 0, 5, 0, 0, 0, 0, 3, 0, 0, 1, 0, 0, 0, -5, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 3, -5, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 3, 3, 0, 0, 0, 0, 3, 0, -3, 0, 0, 0, 1], 'Zax': [0, 3, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 5, 0, 0, 0, 3, 0, 0, 5, 5, 0, 0, 0, 1], 'Jaluka': [0, 5, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0], 'Sanity': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 5, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0], 'Brix': [1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 1, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, -3, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3], 'sage32': [0, 0, 0, 0, 0, 0, 0, -3, 0, 3, 1, 0, 0, 0, 0, 0, 0, 0, 3, 5, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, -5, 0, 0, 0, 1, 0, 0, 5, 3, 3, 0, 0, 0, 1, 0, 0, 0, 0, 0, 5], 'clipper': [5, 0, 0, 0, 0, 0, 0, 1, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, -3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 5, 0, 3, 0, 0, 0, 0, 0, 0, 0, 5, 5, 0, 0, 0, 0, 0, 3, 5, 5, 0, 0, 0, 0], 'YJAM': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 3, 0, 0, 3, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 5, 5, 0, 5, 0, 0], 'Martin': [1, -3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 3, 0, -3, 0, 0, 0, -3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, -3, -3, 0, 0, 0, 5, 0, 0, 0, 0, 0, 1, 3, 5, 0, -3, 0, 0], 'Mike': [5, 0, 0, 0, 0, 0, 5, 0, 0, 1, 0, 5, 3, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 3, 0, 5, 0, 0, 0, 0, 5, 5, 5, 3, 0, 0, 0, 3, 0, 0, 0, 5, 3, 0, 0, 0, 0, 5, 0, 5, 3, 0, 0, 0, 0], 'Strongbad': [5, 0, -3, 0, -3, 0, 0, -5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, -5, 0, 0, 0, 0, 0, 0, 0, 3, 0, 3, 0, 0, 0, -5, 0, 1, 0, 5, 5, 1, 0, 0, 0, 5, 0, 1, 1, 0, 0, 0, 3], 'Brian': [5, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 5, 0, 0, 0, 0, 5, 0, 0, 0, 5, 0, 3, 0, 0, 0, 0, 1, 0, 0, 0, 0, 1, 0, 5, 3, 5, 0, 0, 5, 3, 0, 0, 0, 0, 0, 3, 5, 0, 0, 0, 0], 'Douglas Anderson': [5, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 5, 0, 0, 0, 0, 5, 0, 0, 0, 5, 0, 3, 0, 0, 0, 0, 1, 0, 0, 0, 0, 1, 0, 5, 3, 5, 0, 0, 5, 3, 0, 0, 0, 0, 0, 3, 5, 0, 0, 0, 0], 'chrispeh': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, -5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0], 'Rudy A': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 1, 0, 0, 0, 0, 0, 5, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0], 'McLovin': [0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, -5, 3, 0, -5, 0, 0, 0, 0, 0, 0, 0, 0, -3, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0], 'Mike Williams': [3, 0, 5, 0, 0, 0, 5, 0, 0, 5, 0, 5, 0, 5, 0, 0, 3, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 1, 0, 5, 0, 3, 0, 0, 5, 5, 0, 5, 5, 3, 0, 0, 0, 0, 0, 1, -3, 0, 0, 0, 0], 'Dude': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 3, 0, 0, 5, 0, 0, 0, -3, 0, 0, 0, 0], 'Hideo': [0, 0, -3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 3, 0, -5, 0, 0, 0, 0, 5, 0, 0, 0, 5, 5, 0, 0, 0, 0, 0, 5, 5, 0, 0, 0, 0], 'Rudy Ann': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, -3, 0, 0, -5, 0, 0, 0, 0, 0, 1, 0, 0, 3, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0], 'Don Wang': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 1, 0, 0, 0, 0, 3, 0, 0, 0, 3, 3, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0], 'Anna-Carol': [0, 0, 3, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 3, 0, 0, 0, 5, 0, 0, 0, 3, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0], 'SubRia': [0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 5, 0, 3, -5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 3, 3, -3, 0, 0, 0, 0, 3, 3, 1, -3, 0, 0, 0, 0, 0, 0, 0, 3, 3, 0, 0, 0, 0], 'Ella': [0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 3, 1, 0, 0, 0, 0, 3, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, -5, 1, 0, -5, 0, 5, 0, 0, 1, 0, 0, 3, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0], 'Leslie': [5, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 5, -5, 0, 0, 0, 0, 5, 0, 0, 0, -5, 5, 0, 0, 0, 0, 0, 5, 5, 0, 0, 0, 0], 'Sam': [0, 3, 0, 0, 0, 1, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, -3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0], 'Cedric Jones': [0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 5, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 5, -5, 0, 0, 0, 0, 0, 3, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0], 'McLean': [0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 5, 0, 0, 3, 0, 0, 0, 5, 5, 0, 0, 0, 5, 5, 0, 0, 0, 0, 0, 0, 5, 0, 1, 0, 5], 'Shannon': [3, 0, 0, 0, 0, 0, 0, 0, 0, 5, 5, 0, 0, 0, 0, 0, 5, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 5, 5, -5, 0, 5, 0, 0, 5, 0, 5, 5, 5, 5, 0, 0, 5, 0, 0, 5, 5, 0, 5, 0, 0], 'Joshua': [0, 0, 0, 0, 0, 0, 0, 1, 0, 0, -3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, -5, 0, -5, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0], 'Jeremy': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, -5, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0], 'Pablo Sanchez': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 3, 0, 0, 0, 0, 1, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0], 'Scream': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, -5, 0, -5, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0], 'ROFLOL': [5, 0, 0, 0, 0, 0, 0, 0, 0, 3, 3, 3, 0, 0, 0, 0, 5, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 5, 5, -5, 0, 5, 0, 0, 5, 5, 5, 5, 5, 5, 0, 0, 0, 0, 0, 5, 5, 0, 0, 0, 0], 'Albus Dumbledore': [0, 0, 0, 0, 0, 0, 0, 1, 0, 0, -3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, -3, 0, -5, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0], 'Alexandra': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, -5, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0], 'NaRwHaLs': [5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 3, 0, -5, 0, 0, 5, 0, 5, 0, 0, 0, -3, 3, 0, 0, 0, 0, 0, 5, 5, 0, 0, 0, 0], 'J-Rok': [5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, -3, 0, -5, 0, 0, 0, 0, 0, 0, -5, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0], 'snoosh': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 5, 0, 1, 0, 0, 0, 0, 3, 5, 0, -3, 5, 3, 0, 0, -5, 0, 0, 5, 5, 0, 0, 0, 0], 'Huxil': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 1, 0, 0], 'hidan': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 5, 0, 0, 3, 0, 5, 0, 0, 0, 0, 5, 0, 0, 0, 5, 5, 5, 0, 0, 0, 0, 5, 5, 0, 0, 0, 0], 'Leah': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0], 'Boxxy': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 5, 0, 0, 5, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0], 'Rosanna': [0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 5, 3, 0, 0, -3, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 1, 5, 0, 1, 0, 5, 0, 0, 0, 0, 0, 0, 5, 5, 0, 0, 0, 0, 0, 0, -3, 0, 0, 0, 0], 'Claire': [5, 3, 0, 0, 0, 0, 0, 5, 0, 3, 1, 0, 0, 0, 1, 0, 1, 0, 3, 0, 0, 0, 0, -3, 0, 5, 0, 0, 0, 0, 5, 5, 0, 1, 0, -5, 5, 0, 3, 3, 0, 5, 5, 5, 0, 0, 0, 0, 0, 5, 5, 0, 1, 0, 1], 'David': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, -5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 5, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5], 'Zachary': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 1, 0, 5, 5, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0], 'ParseClaws': [5, 5, 3, 0, 0, 0, 0, 0, 0, -5, 0, 3, 0, 0, 3, 0, 0, 0, 3, 0, 0, 0, 0, 5, 0, 5, 0, 1, 1, 0, 3, 1, 0, 0, 0, -5, 0, 5, 0, 0, 0, 0, 3, 3, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, -3], 'Tiffany': [3, 0, 0, 1, 0, 0, 0, -3, 0, 5, -3, 0, 0, 0, 3, 0, 3, 0, 1, 3, 0, 0, 0, 0, 0, 0, 0, -5, 1, 0, 5, 5, 0, 0, 0, -5, 0, 0, 1, 3, 5, 3, 5, 5, 0, 0, -3, 0, 0, 5, 5, 3, 5, 0, 3], 'Sidney': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 3, 0, 3, 0, 0, 0, 0, 1, 0, 0, 0, 5, 5, 0, 0, 3, 0, 3, 0, 0, 0, 0, 5, 5, 0, 0, 5, 0, 3, 0, 3, 0, 0, 5, 5, 0, 0, 0, 0], 'Tony': [-5, -5, -5, -5, -5, -5, -5, 1, 1, -5, -5, 1, 1, 3, 5, -5, 1, 1, 5, 1, -5, -5, -5, 1, 1, -5, 1, -5, 1, 1, 3, 0, 0, 0, 0, 0, 3, 0, 0, 0, 3, 0, 0, -3, 0, -5, -5, -5, 0, 5, 5, 1, 1, 1, 1], 'Mark': [5, 3, 0, 0, 0, 0, 0, 0, 0, 3, 0, 5, 0, 0, 3, 0, 0, 0, 1, 0, 0, 0, 0, 3, 0, 0, 0, 3, 3, 0, 5, 1, 0, 3, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 5], 'NuNu': [0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 3, 0, 0, 1, 0, 5, 0, 0, 0, 0, 0, 1, 3, 3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0], 'Hamza ALi': [0, 0, 0, 0, 0, 0, 0, -5, 0, 5, 1, 0, 0, 0, 1, 0, 5, 0, -3, 0, 3, 0, 0, 0, 0, 0, 0, -5, 0, 0, -3, 5, 0, 5, 0, 3, 0, 0, 1, 0, 3, 1, 3, 5, 0, 0, 0, 0, 0, 1, -5, 0, 0, 0, 0], 'dencellia': [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, -3, 0, 0, 0, 0, 0, 0, 0, 0, 1, 3, 0, 1, 0, 0, 0, 0, 3, 0, 0, 0, 0, 0, 0, 5, 0, 3, 0, 0, 0, 0, -3, 3, 0, 0, 0, 0], 'Nathan': [0, 0, 0, 0, 0, 0, 0, -5, 0, 5, 1, 0, 0, 0, 1, 0, 5, 0, -3, 0, 3, 0, 0, 0, 0, 0, 0, -5, 0, 0, -3, 5, 0, 5, 0, 3, 0, 0, 1, 0, 3, 1, 3, 5, 0, 0, 0, 0, 0, 1, -5, 0, 0, 0, 0], 'James': [3, 1, 0, 0, 0, 0, 3, -3, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 1, 1, 0, 0, 0, -5, 0, 0, 0, -5, -5, 0, -3, 5, 0, 3, 0, -5, 0, 0, 0, 5, 0, 5, 1, 3, 0, 0, 0, 0, 0, -5, 5, 0, -3, 0, 5], 'KeeLed': [0, 0, 0, 5, 0, 0, 0, 5, 0, 5, 5, 0, 0, 0, 0, 0, 5, -3, -3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 5, 0, 0, 5, 0, 5, 5, 5, 5, 0, 0, 0, 0, 0, 3, 1, 0, 0, 0, 0], 'Megan': [5, 5, 0, 0, 0, 0, 0, 0, 0, 3, 0, 5, 0, 0, 1, 0, 5, 0, 1, 5, 0, 0, 0, 0, 0, 1, 0, 5, 0, 0, 3, 5, 5, 0, 0, 5, 0, 0, 3, 0, 0, 3, 5, 5, 0, 0, 0, 0, 0, 5, 5, 0, 5, 0, 0]})
finally: sys.stdout = sys._jupyter_stdout 
print("Test passed")
```

Output:
```
~\anaconda3\envs\env1\lib\unittest\case.py in assertDictEqual(self, d1, d2, msg)
   1119                            pprint.pformat(d2).splitlines())))
   1120             standardMsg = self._truncateMessage(standardMsg, diff)
-> 1121             self.fail(self._formatMessage(msg, standardMsg))
   1122 
   1123     def assertDictContainsSubset(self, subset, dictionary, msg=None):

~\anaconda3\envs\env1\lib\unittest\case.py in fail(self, msg)
    668     def fail(self, msg=None):
    669         """Fail immediately, with the given message."""
--> 670         raise self.failureException(msg)
    671 
    672     def assertFalse(self, expr, msg=None):

AssertionError: {'Ben[15213 chars]0, 0], 'KittyP': [0, 1, -5, 0, 0, 0, 0, 0, 0, [133 chars], 0]} != {'Ben[15213 chars]0, 0]}
Diff is 82080 characters long. Set self.maxDiff to None to see it.
```

<!-- grade_id: task2-autograder2, points: 1, grade: true, solution: false -->
```python
# Test your code -- check that it works with a file other than ratings.txt
sys._jupyter_stdout = sys.stdout
sys.stdout = open(os.devnull, 'w')    
with open('mtf.txt', 'w') as fileObj: fileObj.write("a, 0 0 1 2\nb, 1 1 5 5")
try: assert_equal(readRatings("mtf.txt"), {'a': [0, 0, 1, 2], 'b': [1, 1, 5, 5]})
finally: sys.stdout = sys._jupyter_stdout 
if os.path.exists("mtf.txt"): os.remove("mtf.txt")
print("Test passed")
```

Output:
```
Test passed
```

### Test `readRatings()`

Write a small function that tests your code by calling `readRatings()` and then printing out each key, value pair in the returned dictionary. 

Check that this is consistent with ratings.txt and your expectations. Once you are satisfied it does, you can comment out the line calling `testReadRatings()` so that this code doesn't run every time you need to re-run the entire file. You do not need to worry about making the output pretty (i.e., you can just let Python handle printing the list). Your output should look like this for the first and last two lines: 

    Ben --> [5, 0, 0, 0, 0, 0, 0, 1, 0, 1, -3, 5, 0, 0, 0, 5, 5, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 0, 0, 1, 3, 0, 1, 0, -5, 0, 0, 5, 5, 0, 5, 5, 5, 0, 5, 5, 0, 0, 0, 5, 5, 5, 5, -5]
    Moose --> [5, 5, 0, 0, 0, 0, 3, 0, 0, 1, 0, 5, 3, 0, 5, 0, 3, 3, 5, 0, 0, 0, 0, 0, 5, 0, 0, 0, 0, 0, 3, 5, 0, 0, 0, 0, 0, 5, -3, 0, 0, 0, 5, 0, 0, 0, 0, 0, 0, 5, 5, 0, 3, 0, 0]
    
    ...
    
    KeeLed --> [0, 0, 0, 5, 0, 0, 0, 5, 0, 5, 5, 0, 0, 0, 0, 0, 5, -3, -3, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 5, 0, 0, 5, 0, 5, 5, 5, 5, 0, 0, 0, 0, 0, 3, 1, 0, 0, 0, 0]
    Megan --> [5, 5, 0, 0, 0, 0, 0, 0, 0, 3, 0, 5, 0, 0, 1, 0, 5, 0, 1, 5, 0, 0, 0, 0, 0, 1, 0, 5, 0, 0, 3, 5, 5, 0, 0, 5, 0, 0, 3, 0, 0, 3, 5, 5, 0, 0, 0, 0, 0, 5, 5, 0, 5, 0, 0]


<!-- grade_id: task2t-answer, points: 0, grade: false, solution: true -->
```python
def testReadRatings():
    # YOUR CODE HERE
    readRatings("ratings.txt")
    for i in customerDB.items():
        print(i)
        
#testReadRatings() #Comment out this line when you are done testing 
```

<!-- grade_id: task2t-autograder, points: 1, grade: true, solution: false -->
```python
# Test your code
sys._jupyter_stdout = sys.stdout
sys.stdout = open(os.devnull, 'w')    
old = readRatings
del readRatings
try: testReadRatings()
except NameError: pass
else: 
    raise AssertionError("testReadRatings() does not call readRatings()")
finally:
    sys.stdout = sys._jupyter_stdout
    readRatings = old
    del old
print("Test passed")
```

Output:
```
Test passed
```

## Task 3: Write `writeRatings()`

Now, you will write a function `writeRatings()` that takes two input parameters, a customer database and a filename, and writes out the contents of the database to the inputted filename.  

### Pseudocode

* Open the file specified by the input parameter `filename` for writing 
* Loop over each key, value pair in the customer database 
    * Convert the ratings list to a string (use a loop and an accumulator to build up this string!)
    * Combine the ratings list string with the customer ID, matching the format in ratings.txt
    * Write the line to the file

<!-- grade_id: task3-answer, points: 0, grade: false, solution: true -->
```python
def writeRatings(customerDB, filename):
    # YOUR CODE HERE
    # open file for writing
    file3 = open(filename, "w")
    # create list for lines
    lines= []
    # create string for line
    line_string = ""
    # iterate over customerDB
    for key, value in customerDB.items():
        # create string for ratings
        ratings_String = ""
        # iterate over vales in customer dict
        for i in value:
            # transform rating from int into string
            str_i = str(i)
            # add rating to string for ratings
            ratings_String += str_i
            # add space between ratings
            ratings_String += " "
        # create line tuple with ussername, a comma, the string of ratings, and \n to start a new line
        line = key+", ", ratings_String+"\n"
        # iterate over tuple elments
        for i in line:
            # add them to line string
            line_string += i
    # add line string to list of lines
    lines.append(line_string)
    # write list of lines as lines in file
    file3.writelines(lines)
    # close file
    file3.close()
    
          
          

```

<!-- grade_id: task3-autograder, points: 3, grade: true, solution: false -->
```python
# Test your code
# Note this test might fail frivolously due to small differences in whitespace
# My solution has a white space at the end of each line and a blank line at the end of each file
# (which is the easiest way to handle writing out the file)
sys._jupyter_stdout = sys.stdout
sys.stdout = open(os.devnull, 'w')    
trdb = {'x': [0, 0, 1, 2], 'y': [1, 1, 5, 5]}
writeRatings(trdb, 'trdb.txt')
with open('trdb.txt') as fo: output = fo.read()
try: assert_equal(output, "x, 0 0 1 2 \ny, 1 1 5 5 \n")
finally: sys.stdout = sys._jupyter_stdout 
if os.path.exists("trdb.txt"): os.remove("trdb.txt")
print("Test passed")
```

Output:
```
Test passed
```

### Test `writeRatings()`

Write a small function that tests your code by calling `readRatings()` and then `writeRatings()`. Your code should read in the real <tt>ratings.txt</tt> file but write out to a dummy file (e.g., <tt>testratings.txt</tt>). This way you can check that your code is able to recreate the same file without overwriting the original ratings file (which would be a very normal thing to do at least once before figuring out the solution!) 

Check that the contents of output file match <tt>ratings.txt</tt>. Once you are satisfied it does, you can comment out the line calling `testWriteRatings()` so that this code doesn't run every time you need to re-run the entire file. 

<!-- grade_id: task3t-answer, points: 0, grade: false, solution: true -->
```python
def testWriteRatings():
    # YOUR CODE HERE
    readRatings("ratings.txt")
    writeRatings(customerDB, "testratings.txt")

#testWriteRatings() #Comment out this line when you are done testing 
```

<!-- grade_id: task3t-autograder, points: 1, grade: true, solution: false -->
```python
# Test your code
sys._jupyter_stdout = sys.stdout
sys.stdout = open(os.devnull, 'w')    
old = writeRatings
del writeRatings
try: testWriteRatings()
except NameError: pass
else: 
    raise AssertionError("testWriteRatings() does not call writeRatings()")
finally:
    sys.stdout = sys._jupyter_stdout
    writeRatings = old
    del old
print("Test passed")
```

Output:
```
Test passed
```

***

# Part 2: A More Sophisticate Recommendation

In this part, you will write a function `sophisticatedRecommendation()` that makes a recommendation to the customer by finding the customer most like themselves, and recommending the book that person rates most highly (and that the customer hasn't already read). 

Again, to help break the task down into more manageable chunks, you will first build a couple of helper functions:
* `computeSimScores()` returns a list of the similarity score between the customer and all other customers in the customer database. 
* `sortUserRatings()` returns a list of tuples containing a book rating, that book's ISBN, and that book's index, sorted from highest to lowest book rating.

## Task 4: Write `computeSimScores()`

`computeSimScores()` should take two input parameters, a target customer and a customerDB, and return a list of all the similarity scores between the customer and all other customers (not including themself). This list should contain tuples of consisting of a similarity score between the target customer and another customer and the "other" customer with which it was computed. 


### Pseudocode

* Initialize an empty list to store the resulting similarity scores
* Get the target customer's list of ratings
* For each customer in the database, get their username and list of ratings
    * If the username doesn't match the target customer
        * initialize their similarity score to 0
        * For each book (looping over the valid book indices)
            * Multiply the target customer's book rating with the "other" customer's book rating and this to their similarity score
        * Append a tuple of their similarity score and the other customer's username to the similarity score list
* Sort and reverse the list of similarity scores so that the entries are ordered from best to worst match


<!-- grade_id: task4-answer, points: 0, grade: false, solution: true -->
```python
def computeSimScores(targetCustomer, customerDB):
    # made simScores global so that it can be called in test function
    global simScores
    simScores = []
    # YOUR CODE HERE
    #create list of ratings for target customer
    tcRating = customerDB.get(targetCustomer)
    #iterate over items in customer dict
    for key, value in customerDB.items():
        #start index counter from -1
        index = -1
        if key != targetCustomer:
            # if the dict key is not the target customer
            # intialize similarity Score
            simScore = 0
            # iterate over ratings
            for i in value:
                #add 1 to index counter
                index +=1
                # multiply other user rating with target costumer rating
                score = i * tcRating[index]
                # add that to simScore
                simScore +=score
            # create tuple for user with their username and simScore
            userSimScore = simScore, key
            # add tuple to user sim score list
            simScores.append(userSimScore)
    # sor list
    simScores.sort()
    # reverse it so larger sim scores are first
    simScores.reverse()
    # return sim score
    return simScores
```

<!-- grade_id: task4-autograder1, points: 2, grade: true, solution: false -->
```python
# Test your code -- Test on data from ratings.txt
sys._jupyter_stdout = sys.stdout
sys.stdout = open(os.devnull, 'w')
customers = readRatings("ratings.txt")
try: assert_equal(computeSimScores("Mike Williams", customers), [(181, 'Cust6'), (178, 'Cust8'), (160, 'ROFLOL'), (156, 'Megan'), (154, 'Ben'), (142, 'Francois'), (142, 'Cust1'), (137, 'KeeLed'), (131, 'Mike'), (127, 'Claire'), (124, 'Shannon'), (123, 'Nathan'), (123, 'Hamza ALi'), (122, 'andrew'), (118, 'joe'), (118, 'Rosanna'), (106, 'matt c'), (106, 'Tiffany'), (104, 'McLean'), (102, 'Ocelot'), (101, 'Sidney'), (96, 'Douglas Anderson'), (96, 'Brian'), (94, 'Mark'), (92, 'Moose'), (85, 'Cust7'), (83, 'hidan'), (76, 'ender'), (74, 'Strongbad'), (73, 'James'), (70, 'NuNu'), (67, 'clipper'), (65, 'Priscilla'), (62, 'snoosh'), (62, 'Iren'), (62, 'Cust5'), (60, 'Sanity'), (56, 'Apollo'), (55, 'Boxxy'), (54, 'Jim C'), (52, 'Cust10'), (51, 'sage32'), (49, 'Anna-Carol'), (48, 'Rudy A'), (47, 'SubRia'), (47, 'I. Ned'), (46, 'Cust9'), (45, 'ma'), (45, 'Zax'), (45, 'Dude'), (42, 'ParseClaws'), (36, 'Don Wang'), (36, 'Cedric Jones'), (35, 'Ella'), (35, 'Cust3'), (33, 'Pablo Sanchez'), (27, 'NaRwHaLs'), (25, 'Leslie'), (25, 'Cust2'), (24, 'YJAM'), (23, 'Zachary'), (22, 'Reuven'), (20, 'David'), (19, 'Hideo'), (14, 'dencellia'), (14, 'crom'), (11, 'Harry Potter'), (10, 'Jaluka'), (7, 'Martin'), (3, 'Brix'), (1, 'Sam'), (0, 'Leah'), (-1, 'Cust4'), (-2, 'Albus Dumbledore'), (-5, 'Rudy Ann'), (-7, 'Jeremy'), (-7, 'Alexandra'), (-8, 'Huxil'), (-9, 'McLovin'), (-12, 'Joshua'), (-22, 'chrispeh'), (-25, 'Scream'), (-40, 'J-Rok'), (-68, 'Tony')])
finally: sys.stdout = sys._jupyter_stdout 
print("Test passed")
```

Output:
```
~\anaconda3\envs\env1\lib\unittest\case.py in assertSequenceEqual(self, seq1, seq2, msg, seq_type)
   1008         standardMsg = self._truncateMessage(standardMsg, diffMsg)
   1009         msg = self._formatMessage(msg, standardMsg)
-> 1010         self.fail(msg)
   1011 
   1012     def _truncateMessage(self, message, diff):

~\anaconda3\envs\env1\lib\unittest\case.py in fail(self, msg)
    668     def fail(self, msg=None):
    669         """Fail immediately, with the given message."""
--> 670         raise self.failureException(msg)
    671 
    672     def assertFalse(self, expr, msg=None):

AssertionError: Lists differ: [(181[1334 chars]-25, 'Scream'), (-25, 'KittyP'), (-40, 'J-Rok'), (-68, 'Tony')] != [(181[1334 chars]-25, 'Scream'), (-40, 'J-Rok'), (-68, 'Tony')]

First differing element 82:
(-25, 'KittyP')
(-40, 'J-Rok')

First list contains 1 additional elements.
First extra element 84:
(-68, 'Tony')

Diff is 1657 characters long. Set self.maxDiff to None to see it.
```

<!-- grade_id: task4-autograder2, points: 1, grade: true, solution: false -->
```python
# Test your code -- test on other data
sys._jupyter_stdout = sys.stdout
sys.stdout = open(os.devnull, 'w')
customers = readRatings("ratings.txt")
try: assert_equal(computeSimScores('x', {'x': [0, 0, 1, 2], 'y': [1, 1, 5, 5]}), [(15, 'y')])
finally: sys.stdout = sys._jupyter_stdout 
print("Test passed")
```

Output:
```
Test passed
```

### Test `computeSimScores()`

Write a small function that tests your code. You will need to call `readRatings()` to initialize the databases. Call `computeSimScore()` on a couple of existing customers to test the result. You may want to calculate the similarity scores for one customer in excel or create a small test customer database to support your testing. 

When you are convinced your code works, comment out the line calling your test function. 

<!-- grade_id: task4t-answer, points: 0, grade: false, solution: true -->
```python
#Solution
def testComputeSimScores():
    # YOUR CODE HERE
    readRatings("ratings.txt")
    computeSimScores("Mike Williams", customerDB)
    print(simScores)
# testComputeSimScores() #Comment out this line when you are done testing 
```

Output:
```
[(181, 'Cust6'), (178, 'Cust8'), (160, 'ROFLOL'), (156, 'Megan'), (154, 'Ben'), (142, 'Francois'), (142, 'Cust1'), (137, 'KeeLed'), (131, 'Mike'), (127, 'Claire'), (124, 'Shannon'), (123, 'Nathan'), (123, 'Hamza ALi'), (122, 'andrew'), (118, 'joe'), (118, 'Rosanna'), (106, 'matt c'), (106, 'Tiffany'), (104, 'McLean'), (102, 'Ocelot'), (101, 'Sidney'), (96, 'Douglas Anderson'), (96, 'Brian'), (94, 'Mark'), (92, 'Moose'), (85, 'Cust7'), (83, 'hidan'), (76, 'ender'), (74, 'Strongbad'), (73, 'James'), (70, 'NuNu'), (67, 'clipper'), (65, 'Priscilla'), (62, 'snoosh'), (62, 'Iren'), (62, 'Cust5'), (60, 'Sanity'), (56, 'Apollo'), (55, 'Boxxy'), (54, 'Jim C'), (52, 'Cust10'), (51, 'sage32'), (49, 'Anna-Carol'), (48, 'Rudy A'), (47, 'SubRia'), (47, 'I. Ned'), (46, 'Cust9'), (45, 'ma'), (45, 'Zax'), (45, 'Dude'), (42, 'ParseClaws'), (36, 'Don Wang'), (36, 'Cedric Jones'), (35, 'Ella'), (35, 'Cust3'), (33, 'Pablo Sanchez'), (27, 'NaRwHaLs'), (25, 'Leslie'), (25, 'Cust2'), (24, 'YJAM'), (23, 'Zacha
```

<!-- grade_id: task4t-autograder, points: 1, grade: true, solution: false -->
```python
# Test your code
sys._jupyter_stdout = sys.stdout
sys.stdout = open(os.devnull, 'w')    
old = computeSimScores
del computeSimScores
try: testComputeSimScores()
except NameError: pass
else: 
    raise AssertionError("testComputeSimScores() does not call computeSimScores()")
finally:
    sys.stdout = sys._jupyter_stdout
    computeSimScores = old
    del old
print("Test passed")
```

Output:
```
Test passed
```


## Task 5: Write `sortUserRatings()`

The customer database stores for each customer a list of their book ratings, order by book index. This has worked well throughout the assignment, but in this stage using the order of the items in the list to represent the book's index is a bit complicated. We *could* loop through the list to find the highest score, see if the customer has read that book, and if yes, loop again to find the second highest score, etc. But that is very inefficient, and still quite complicated to write. 

What we would really like to do is take the list of ratings from the
target customers best match, and then sort their book ratings so that we can easily access the top ratings in order, until we find one that the target customer hasn't read. 

But if we sort the list of ratings, we lose the ability to infer which rating goes with which book. So we will need to make a new data structure that keeps track of the rating and data indicating which book it belongs to. To further support our future needs we might want to keep track of both the ISBN number (which can access the book from the book database) and the index (which can access the book from a list of ratings).  

Putting that all together: `sortUserRatings()` should take two input parameters, a list of book ratings, and a bookDB. It will produce a list of tuples, in which the first element is a rating, the second is the book's ISBN, and the third, is the book's index. 


### Pseudocode

* Initialize an empty list to store the resulting list of customer ratings tuples
* Loop over each book in the book database
    * Get the index of the book from the book's data
    * Use that to access the rating from the rating list
    * Append a tuple containing the rating, the book's ISBN number, and the book's index to the list of customer ratings tuples
* Sort the list of customer ratings tuples
* Reverse the sort of that list
* Return the list


<!-- grade_id: task5-answer, points: 0, grade: false, solution: true -->
```python
def sortUserRatings(u_ratings, bookDB):
    # YOUR CODE HERE
    # made userRatings global so it can be called in test function
    global userRatings
    # intialize userRatings list
    userRatings = []
    #iterate over bookDB items
    for key, value in bookDB.items():
        # get index form fourth value
        index = value[3]
        #get isbn from key
        isbn = key
        # make tuple with that user's ratings, and the books isbn and index
        userBookRatings = u_ratings[index], isbn, index
        # add tuple to list of userRatings
        userRatings.append(userBookRatings)
    # sort ratings
    userRatings.sort()
    # reverse them so that higher rated books go first
    userRatings.reverse()
    # return list of userRatings
    return userRatings

```

<!-- grade_id: task5-autograder1, points: 2, grade: true, solution: false -->
```python
# Test your code -- test on data from ratings.txt, books.txt
sys._jupyter_stdout = sys.stdout
sys.stdout = open(os.devnull, 'w')
customers = readRatings("ratings.txt")
books = readBooks("books.txt")
try: assert_equal(sortUserRatings(customers["Mike Williams"], books), [(5, '8495618818', 9), (5, '3641082505', 11), (5, '0937247065', 2), (5, '0590353403', 42), (5, '0553382570', 6), (5, '0547416113', 33), (5, '0545229936', 13), (5, '0375890369', 38), (5, '0375846724', 41), (5, '0330338730', 39), (3, '8434878607', 43), (3, '1905654391', 35), (3, '0606338039', 16), (3, '0517693119', 0), (3, '0399501487', 18), (1, '0871299208', 30), (1, '0261102958', 49), (1, '0064404990', 31), (0, '8809744454', 3), (0, '1894448200', 21), (0, '1863330550', 20), (0, '1603034013', 52), (0, '1595141712', 5), (0, '1591164418', 28), (0, '1569319006', 27), (0, '1443102067', 46), (0, '1416549178', 40), (0, '1411469380', 37), (0, '1156393361', 44), (0, '0929631048', 4), (0, '0871295377', 25), (0, '0743273565', 14), (0, '0731814909', 8), (0, '0618439102', 26), (0, '0547538648', 53), (0, '0441007465', 17), (0, '0440240167', 34), (0, '0425190374', 22), (0, '0425116840', 12), (0, '0394747232', 47), (0, '0375890291', 7), (0, '0345439740', 19), (0, '0312331754', 32), (0, '0230738034', 29), (0, '0230037518', 10), (0, '0143181538', 51), (0, '0143176763', 24), (0, '0143038095', 48), (0, '0142411736', 15), (0, '0142300942', 45), (0, '0141038462', 54), (0, '0061975133', 36), (0, '0060850523', 23), (0, '0025002600', 1), (-3, '0345339683', 50)])
finally: sys.stdout = sys._jupyter_stdout 
print("Test passed")
```

Output:
```
Test passed
```

<!-- grade_id: task5-autograder2, points: 1, grade: true, solution: false -->
```python
# Test your code -- test on data not from ratings.txt, books.txt
sys._jupyter_stdout = sys.stdout
sys.stdout = open(os.devnull, 'w')
customers = {'x': [0, 0, 1, 2], 'y': [1, 1, 5, 5]}
books = {'01': ('1900', 'abc', 'def', 0), '02': ('1901', 'ghi', 'klm', 1), '03': ('1902', 'nop', 'qrs', 2), '04': ('1903', 'tuv', 'wxy', 3)}
try: assert_equal(sortUserRatings(customers["x"], books), [(2, '04', 3), (1, '03', 2), (0, '02', 1), (0, '01', 0)])
finally: sys.stdout = sys._jupyter_stdout 
print("Test passed")
```

Output:
```
Test passed
```

### Test `sortUserRatings()`

Write a small function that tests your code. You will need to call `readRatings()` and `readBooks()` to initialize the databases. Call `sortUserRatings()` on a couple of existing customers to test the result. You should be able to tell if your code is working by spot testing a few of the results (e.g., find the customer's top rating in ratings.txt and check that this is correct). (Don't worry about how ties should be ordered).

When you are convinced your code works, comment out the line calling your test function. 

<!-- grade_id: task5t-answer, points: 0, grade: false, solution: true -->
```python
def testSortUserRatings():
    # YOUR CODE HERE
    readRatings("ratings.txt")
    readBooks("books.txt")
    sortUserRatings(customerDB["Mike"], bookDB)
    print(userRatings)
    sortUserRatings(customerDB["Priscilla"], bookDB)
    print(userRatings)
    sortUserRatings(customerDB["Ocelot"], bookDB)
    print(userRatings)
    sortUserRatings(customerDB["Mike Williams"], bookDB)
    print(userRatings)
    sortUserRatings(customerDB["Sanity"], bookDB)
    print(userRatings)
    
#testSortUserRatings() #Comment out this line when you are done testing 
```

Output:
```
[(5, '3641082505', 11), (5, '0871299208', 30), (5, '0871295377', 25), (5, '0553382570', 6), (5, '0517693119', 0), (5, '0394747232', 47), (5, '0375846724', 41), (5, '0312331754', 32), (5, '0261102958', 49), (5, '0064404990', 31), (3, '1411469380', 37), (3, '0590353403', 42), (3, '0547416113', 33), (3, '0441007465', 17), (3, '0425116840', 12), (3, '0345339683', 50), (3, '0060850523', 23), (1, '8495618818', 9), (0, '8809744454', 3), (0, '8434878607', 43), (0, '1905654391', 35), (0, '1894448200', 21), (0, '1863330550', 20), (0, '1603034013', 52), (0, '1595141712', 5), (0, '1591164418', 28), (0, '1569319006', 27), (0, '1443102067', 46), (0, '1416549178', 40), (0, '1156393361', 44), (0, '0937247065', 2), (0, '0929631048', 4), (0, '0743273565', 14), (0, '0731814909', 8), (0, '0618439102', 26), (0, '0606338039', 16), (0, '0547538648', 53), (0, '0545229936', 13), (0, '0440240167', 34), (0, '0425190374', 22), (0, '0399501487', 18), (0, '0375890369', 38), (0, '0375890291', 7), (0, '0345439740', 1
```

<!-- grade_id: task5t-autograder, points: 1, grade: true, solution: false -->
```python
# Test your code
sys._jupyter_stdout = sys.stdout
sys.stdout = open(os.devnull, 'w')    
old = sortUserRatings
del sortUserRatings
try: testSortUserRatings()
except NameError: pass
else: 
    raise AssertionError("testSortUserRatings() does not call sortUserRatings()")
finally:
    sys.stdout = sys._jupyter_stdout
    sortUserRatings = old
    del old
print("Test passed")
```

Output:
```
Test passed
```


## Task 6: Write `sophisticatedRecommendation()`

`sophisticatedRecommendation()` should take three input parameters, a target customer, a customerDB, and a bookDB, and recommend a book (that the customer hasn't already read) by finding the most similar other customer, and finding the highest rated book in their list of ratings that the customer hasn't already read. 


### Pseudocode

* Call `computeSimScores()` to get a list of similarity scores between the target customer and all other customers
* The best match will be the first one in the list of similarity scores
* Get a list of possible recommendations by calling sortUserRatings, passing the match's rating list and the book DB
* Loop over each element of the list of possible recommendations
    * Use the possible recommendation's index to look up the customer's rating of that book. If it's zero:
        * We have a match! Print out a message (see below) and return
* If you get here, it means you've run out of recommendation possibilities. Apologize to the customer and then return<sup>*</sup>            

Note that for the test cell to work your code must print out output exactly as: 

    You might like: A Great and Terrible Beauty by Libba Bray rated 5 by Cust6
    
That is `"You might like: " + <title> + " by " + <author> + " rated " + <rating> + " by " + <best match>`. Other formatting is acceptable and will be accepted during grading.

<sup>*</sup> *While it would be better to move on to the second best match if the customer has read all the possible recommendations from their best match, it is not necessary for the assignment (and if you really wanted to address this limitation, tackling the even more sophisticated algorithm described on MyCourses would be the better way to go)* 


<!-- grade_id: task6-answer, points: 0, grade: false, solution: true -->
```python
def sophisticatedRecommendation(targetCustomer, customerDB, bookDB):
    # YOUR CODE HERE
    #compute sims socores for target user
    computeSimScores(targetCustomer, customerDB)
    # get best match from first user in simScores
    bestMatch = simScores[0][1]
    # sort best match ratings
    sortUserRatings(customerDB[bestMatch], bookDB)
    # iterate over best match user ratings
    for i in userRatings:
        # get the book index
        matchIndex = i[2]
        # get the book ISBN
        matchISBN = i[1]
        # get the rating
        matchRating = i[0]
        # if the rating for the target customer for that book is 0
        if customerDB[targetCustomer][matchIndex] == 0:
            # print recommendation
            print("You might like:")
            print(" ",bookDB[matchISBN][1], "("+ bookDB[matchISBN][0] +")","by", bookDB[matchISBN][2])
            print("  Rated", matchRating, "by", bestMatch)
            # exit function
            return
    # if it finishes iterating, print message for no recommendations
    print("Sorry, we don't have any recommendations at this moment")
    # exit function
    return
```

<!-- grade_id: task6-autograder, points: 3, grade: true, solution: false -->
```python
# Test your code
customers = readRatings("ratings.txt")
books = readBooks("books.txt")
with mock.patch('sys.stdout', new=io.StringIO()) as fake_stdout:
    sophisticatedRecommendation("Mike Williams", customers, books)    
assert_equal(fake_stdout.getvalue(), 'You might like: A Great and Terrible Beauty by Libba Bray rated 5 by Cust6\n')
print("Test Passed")
```

Output:
```
~\anaconda3\envs\env1\lib\unittest\case.py in assertMultiLineEqual(self, first, second, msg)
   1201             diff = '\n' + ''.join(difflib.ndiff(firstlines, secondlines))
   1202             standardMsg = self._truncateMessage(standardMsg, diff)
-> 1203             self.fail(self._formatMessage(msg, standardMsg))
   1204 
   1205     def assertLess(self, a, b, msg=None):

~\anaconda3\envs\env1\lib\unittest\case.py in fail(self, msg)
    668     def fail(self, msg=None):
    669         """Fail immediately, with the given message."""
--> 670         raise self.failureException(msg)
    671 
    672     def assertFalse(self, expr, msg=None):

AssertionError: 'You might like:\n  A Great and Terrible Beauty (2010) by[29 chars]t6\n' != 'You might like: A Great and Terrible Beauty by Libba Bra[16 chars]t6\n'
+ You might like: A Great and Terrible Beauty by Libba Bray rated 5 by Cust6
- You might like:
-   A Great and Terrible Beauty (2010) by Libba Bray
-   Rated 5 by Cust6
```

### Test `sophisticatedRecommendation()`

Write a small function that tests your code. You will need to call `readRatings()` and `readBooks()` to initalize the databases. Call `sophisticatedRecommendation()` on a couple of existing customers to test the result. Have at least one test case where the perons has read the most highly rated book.

When you are convinced your code works, you can comment out the line calling your test function. 

<!-- grade_id: task6t-answer, points: 0, grade: false, solution: true -->
```python
def testSophisticatedRecommendation():  
    # YOUR CODE HERE
    readRatings("ratings.txt")
    readBooks("books.txt")
    sophisticatedRecommendation("Mike Williams", customerDB, bookDB)

# testSophisticatedRecommendation() #Comment out this line when you are done testing 
```

Output:
```
You might like:
  A Great and Terrible Beauty (2010) by Libba Bray
  Rated 5 by Cust6
```

<!-- grade_id: task6t-autograder, points: 1, grade: true, solution: false -->
```python
# Test your code
sys._jupyter_stdout = sys.stdout
sys.stdout = open(os.devnull, 'w')    
old = sophisticatedRecommendation
del sophisticatedRecommendation
try: testSophisticatedRecommendation()
except NameError: pass
else: 
    raise AssertionError("testSophisticatedRecommendation() does not call sophisticatedRecommendation()")
finally:
    sys.stdout = sys._jupyter_stdout
    sophisticatedRecommendation = old
    del old
print("Test passed")
```

Output:
```
Test passed
```

***
# Part 3: Putting it all Together

We've now built all the individual pieces, but we don't have a way of running them together. In this part you will write a function `runBookRecommender()` that provides a text based interface for calling the different functions. This function will be very similar to `runBookRecommender_basic()` which you wrote for the last assignment, only now, instead of using print statements to describe what would be done, the code should call the functions you've implemented, e.g., opening accounts, rating books, making recommendations, etc. 

## Import Functions From Part I

First, you will need to include definitions to the key functionality written in the last assignment. For this you will need to either ensure that solA2.py has been downloaded from MyCourses and copied into the same directory as this notebook file, or you will need to replace the line `import solA2` below with copies of the following functions definitions from your last assignment: 

    openAccount()
    login()
    queryRating()
    rateBooks()
    getOneBooksRatings()
    getAverageRatings()
    topThree()
    basicRecommendation()
    
    

<!-- grade_id: cell-46c1ba19645a695e, points: 0, grade: true, solution: true -->
```python
# Download the file solA2.py from MyCourses and place in the same directory 
# as this notebook, or replace the import line below with your function definitions for
# openAccount(), login(), queryRating(), rateBooks(), getOneBooksRatings(), 
# getAverageRatings(), topThree(), and basicRecommendation() from Part 1. 

#from solA2 import *

# YOUR CODE HERE
def openAccount(customerDB, numBooks):
    # loop that checks if username is taken and asks for new input if it is taken and ends loop if it is not
    while True:
        username = input("Create your username: ")
        if username in customerDB.keys():
            print(username, "is already in use, please try a different username")
        else:
            break
    # adds new username and ratings of 0 (unrated) to all books
    customerDB[username] = [0]*numBooks
    # prints successful account creation message
    print("Account was created successfully. Please log in to continue")
    return username

def login(customerDB):
    global username
    username =""
    # Ask user for username, in loop to allow for retaking input when username not in database
    while True:
        # If username in database print success message and return username
        username = input("Type in your username: ")
        if username in customerDB.keys():
            print("Login succeeded. \nWelcome back", username)
            return username
        else:
            print("Sorry, this username does not exist")
            # Ask user if they would like to try again
            retry = input("Would you like to try again?\n Press 'Y' for Yes and 'N' for No\n")
            retry = retry.lower()
            # if yes continue loop and ask for username again
            if retry == "y":
                continue
            # if user does not want to try again print error message and return none
            elif retry == "n":
                print("Login failed")
                return None
            else: 
                print("invalid input")
                print("Login failed")
                return None
            
def queryRating(currentScore):
    SCOREMAP = [0, -5, -3, 1, 3, 5]
    # Changed newScore to Global beacuse otherwise the testQueryRating() could not call this variable
    global newScore
    # Loop that presents current rating and takes user_score as user input to change newScore
    while True:
        #If current score is 0, option to delete is not given and input text is slightly different
        if currentScore == 0:
            print("You have not rated this book")
            user_score = input("To rate this book please enter a score (1-5)\nTo leave the score unchanged press 'Enter'\nTo quit press 'q'\n")
            if user_score == "1" or user_score == "2" or user_score == "3" or user_score == "4" or user_score == "5":
                # turn user input into integer
                user_score = int(user_score)
                # transform human readable rating into machine readable
                newScore = SCOREMAP[user_score]
                break
            elif user_score == "q" or user_score == "Q":
                newScore = "q"
                break
            # If no input is given print message and make newScore the same as currentScore
            elif user_score == "":
                print("I do not want to rate this book")
                newScore = currentScore
                break
            else:
                # if other input is given print error message and continue loop
                print("Invalid Input")
                continue
        else:
            # if currentScore is not 0, presents currentScore, and the option to delete
            # gives current rating by transforming currentScore from machine readable format to human readable
            print("You have rated this book", SCOREMAP.index(currentScore))
            user_score = input("To rate this book please enter a score (1-5)\nTo delete your score press 'd'\nTo leave the score unchanged press 'Enter'\nTo quit press 'q'\n")
            if user_score == "1" or user_score == "2" or user_score == "3" or user_score == "4" or user_score == "5":
                 # turn user input into integer
                user_score = int(user_score)
                # transform human readable rating into machine readable
                newScore = SCOREMAP[user_score]
                break
            # sets newScore to 0 beacse user asked to delete rating
            elif user_score == "d" or user_score == "D":
                newScore = 0 
                break
            elif user_score == "q" or user_score == "Q":
                newScore = "q"
                break
            elif user_score == "":
                # If no input is given print message and make newScore the same as currentScore
                print("I want to keep my current rating")
                newScore = currentScore
                break
            else:
                # if other input is given print error message and continue loop
                print("Invalid Input")
                continue
    # After exiting loop, returns the new value for newScore
    return newScore

def rateBooks(username, customerDB, bookDB):
    # iterate through the items in the dictonary
    for key, values in bookDB.items():
        # print book info for customer
        print(" ")
        print(values[1], "("+values[0]+") by", values[2])
        print(" ")
        # call queryRating function using value 3 in that item in bookDB as index for value in  a user's customerDB rating list
        queryRating(customerDB[username][bookDB[key][3]])
        # if newScore is q, exit function
        if newScore == "q":
            return
        else:
            # otherwise update value in a user's customerDB rating using value 3 that item in bookDB as index for list 
            (customerDB[username][bookDB[key][3]]) = newScore
            continue
    return

def getOneBooksRatings (customerDB,bookIndex):
    # made ratings global so that I could call it in the test function
    global OBratings 
    # Initialize an empty list to hold the resulting ratings
    OBratings = []
    # Loop through each username in the customerDB:
    for username in customerDB:
        # get the rating for the book indicated by book index
        ub_rating = customerDB[username][bookIndex]
        # if the rating is not 0 append the rating to the ratings list
        if ub_rating != 0:
            OBratings.append(customerDB[username][bookIndex])
        # else do nothing
        else:
            pass
    # exit function and return ratings list
    return OBratings

def getAverageRatings(customerDB, bookDB): 
    # make averageRatings global so that it can be called later
    global averageRatings
    # initaiate averageRatings list
    averageRatings = []    
    # iterate over items in bookDB
    for key, values in bookDB.items():
        #get the ratings for that book in customerDB by sing the index in bookDB
        bookIndex = values[3]
        getOneBooksRatings(customerDB, bookIndex)
        # if the length of ratings is more that 0 (i.e. at least one person has rated it)
        if len(OBratings) > 0:
            # compute the average ratings
            avg = sum(OBratings)/len(OBratings)
            # append average ratings, and the key for that book to the list averageRatings
            averageRatings.append(tuple((avg, key)))
    # sort list by greatest rating first
    averageRatings.sort(reverse = True)
    # exit function and return list
    return averageRatings

def topThree(customerDB, bookDB): 
    # Call getAverageRatings
    getAverageRatings(customerDB, bookDB)
    # Print header
    print("---------------------------------------------------")
    print("The three highest rated books in our database are. ")
    print("---------------------------------------------------")
    # iterate over first three items
    for i in range(3):
        # iterate over items in bookDB
        for key, value in bookDB.items():
            # if the second value in an item in the list averageRatings matches the dictionary key
            if averageRatings[i][1] == key:
                # print book info for customer
                print(value[1], "("+value[0]+") by", value[2])
                # print that book's average rating rounded
                print("Average Rating: ", (round(averageRatings[i][0], 2)))
                print(" ")

def basicRecommendation(username, customerDB, bookDB):    
    # call getAverageRatings
    getAverageRatings(customerDB, bookDB)
    # Loop through averageRatings list
    for i in range(len(averageRatings)):
        # loop through bookDB
        for key, value in bookDB.items():
            # get ISBN from averageRatings
            if averageRatings[i][1] == key:
                    # get the index fom bookDB and if the rating for that book from csutomerDB and check if it is 0
                if customerDB[username][bookDB[key][3]] == 0:
                # if it is print reccommendation for user and the average rating for the recommended book
                    print("You might like: ")
                    print("  ", value[1], "("+value[0]+") by", value[2])
                    print("   Average Rating: ", (round(averageRatings[i][0], 2)))
                    # exit function
                    return
                # if the user has rated book (i.e. rating in customerDB is not 0) continue loop
                else:
                    pass   
        # if loop finishes print message indicating that there are no recommendations 
        print("We have no recommendations at this time. You have rated all books in the database.")

```

```python
#runBookRecommender kept the previous username saved even after pressing q 
#so I wrote a logout function to set username back to empty
def logout():
    global username
    username = ""
    return username
```

## Task 7: Write `runBookRecommender()`

`runBookRecommender()` takes no input parameters. It should provide a menu that allows the customer to pick the functionality that they want from the choices. When finished with it, present the menu again until the customer chooses to exit. A good strategy for this task would be to copy over your implementation to `runBookRecommender_basic()` and then replace each of the print statements with a call to the appropriate function. 

You will also need to make calls to `readBooks()` and `readRatings()` at the very start to initialize your book and customer databases and you will need to make sure the last thing you do before exiting is call `writeBooks()` to write out any changes (*NB do not try to call writeBooks after each change to the database, this will be too complicated. It is enough to do this after the user indicates that they are ready to quit*). 

Remember that many functions depend on a particular user being 'logged in'. So you will also need to keep track of who (if anyone is logged in) and check (before calling certain functions) that someone is logged in. [A sample run should look something like this.](https://drive.google.com/file/d/17ufCObMgT1OZGLlGBYtZCgKpPZsi9V0J/view?usp=sharing)


    
### Pseudocode
* Create a variable to track the active (i.e., logged in) customer
* Load external data (e.g., `books.txt` and `ratings.txt`)
* Looping until the customer 'quits': 
    * Provide the customer with a list of options: 
    * Depending on their selection, call the desired function. Note: 
        * That for many functions you first check that someone is logged in
        * When the customer quits the customerDB should first be written out to file

<!-- grade_id: task7-answer, points: 3, grade: true, solution: true -->
```python
def runBookRecommender():
    # YOUR CODE HERE
    # intialize global username variable
    global username
    # set it to empty
    username = ""
    readBooks("books.txt")
    readRatings("ratings.txt")
    print("----------------------")
    print(" BOOK RECOMMENDER 5000 ")
    print("----------------------")
    print(" ")
    # create list containing all menu options. I changed top 10 books to top three because that is what makes sense with the other parts of the assignment.
    mainmenu = [" ", "1. Create an account", "2. Log in", "3. Rate books", "4. Display a generic list of the 3 most popular books", "5. Get a generic book recommendation", "6. Get a personalized book recommendation"]
    #start loop that displays menu options and takes user input until q is pressed
    while True:
        # print all menu options
        for i in mainmenu:
            print(i)
        # take user input to select menu option print corresponding message and carry out corresponding function
        option = input("\nEnter an option(1-6), or Q to quit: ")
        if option == "1":
            print("\n*** Creating an account... ***\n")
            openAccount(customerDB, len(bookDB))
        elif option == "2":
            print("\n*** Logging in... ***\n")
            login(customerDB)
        elif option == "3":
            print("\n*** Rating books... ***\n")
            # run function only if username is  not empty, else print error message
            if username:
                rateBooks(username, customerDB, bookDB)
            else:
                print("You must be logged in to rate books")
        elif option == "4":
            print("\n*** Displaying top 3... ***\n")
            topThree(customerDB, bookDB)
        elif option == "5":
            print("\n*** Making a generic recommendation... ***\n")
            # run function only if username is  not empty, else print error message
            if username: 
                basicRecommendation(username, customerDB, bookDB)
            else:
                print("You must be logged in to get recommendations")
        elif option == "6":
            print("\n*** Making a personalized recommendation... ***\n")
            # run function only if username is  not empty, else print error message
            if username: 
                sophisticatedRecommendation(username, customerDB, bookDB)
            else:
                print("You must be logged in to get recommendations")   
        elif option == "q" or option == "Q": 
            # if username is not empty run logout function to set it back to empty
            if username:
                print("\n*** Logging out... ***\n")
                username = logout()
                break
            else:
                print("\n*** Quitting... ***\n")
            break
        # if invalid input (not q or not 1-6) print error message and continue loop
        else:
            print("\n Invalid input. Please try again\n")
    # after loop ends update file and print goodbye message
    writeRatings(customerDB, "ratings.txt")
    print(" ")
    print("-----------------------")
    print(" HAVE A NICE DAY. BYE ")
    print("-----------------------")
    return
    
runBookRecommender()
```

Output:
```
----------------------
 BOOK RECOMMENDER 5000 
----------------------
 
 
1. Create an account
2. Log in
3. Rate books
4. Display a generic list of the 3 most popular books
5. Get a generic book recommendation
6. Get a personalized book recommendation

Enter an option(1-6), or Q to quit: 6

*** Making a personalized recommendation... ***

You must be logged in to get recommendations
 
1. Create an account
2. Log in
3. Rate books
4. Display a generic list of the 3 most popular books
5. Get a generic book recommendation
6. Get a personalized book recommendation

Enter an option(1-6), or Q to quit: q

*** Quitting... ***

 
-----------------------
 HAVE A NICE DAY. BYE 
-----------------------
```

###### Test `runBookRecommender()`

Run the above ceel to call your `runBookRecommender()` function and test that everything is working together properly. Cut and paste the output from your run to demonstrate how you have tested each branch of your overall system. Note that if this part isn't running, then you should include a paragraph or two describing how you would test your code for partial credit. *(Note: After pasting your output below, select all lines and hit `tab` to format your answer as monospaced plain text)*

<!-- grade_id: task7t-answer, points: 1, grade: true, solution: true -->
    ----------------------
    BOOK RECOMMENDER 5000 
    ----------------------


    1. Create an account
    2. Log in
    3. Rate books
    4. Display a generic list of the 3 most popular books
    5. Get a generic book reccommendation
    6. Get a personalized book recommendation

    Enter an option(1-6), or Q to quit: 2

    *** Logging in... ***

    Type in your username: KittyP
    Sorry, this username does not exist
    Would you like to try again?
     Press 'Y' for Yes and 'N' for No
    Y
    Type in your username: KittyP
    Sorry, this username does not exist
    Would you like to try again?
     Press 'Y' for Yes and 'N' for No
    n
    Login failed

    1. Create an account
    2. Log in
    3. Rate books
    4. Display a generic list of the 3 most popular books
    5. Get a generic book reccommendation
    6. Get a personalized book recommendation

    Enter an option(1-6), or Q to quit: 3

    *** Rating books... ***


    The Hitchhiker's Guide To The Galaxy (1989) by Douglas Adams

    You must be logged in to rate books

    1. Create an account
    2. Log in
    3. Rate books
    4. Display a generic list of the 3 most popular books
    5. Get a generic book reccommendation
    6. Get a personalized book recommendation

    Enter an option(1-6), or Q to quit: 4

    *** Displaying top 3... ***

    ---------------------------------------------------
    The three highest rated books in our database are. 
    ---------------------------------------------------
    Sabriel (2009) by Garth Nix
    Average Rating:  4.0

    The Bourne Series (2004) by Robert Ludlum
    Average Rating:  3.86

    The Hitchhiker's Guide To The Galaxy (1989) by Douglas Adams
    Average Rating:  3.83


    1. Create an account
    2. Log in
    3. Rate books
    4. Display a generic list of the 3 most popular books
    5. Get a generic book reccommendation
    6. Get a personalized book recommendation

    Enter an option(1-6), or Q to quit: 5

    *** Making a generic recommendation... ***

    You must be logged in to get recommendations

    1. Create an account
    2. Log in
    3. Rate books
    4. Display a generic list of the 3 most popular books
    5. Get a generic book reccommendation
    6. Get a personalized book recommendation

    Enter an option(1-6), or Q to quit: 6

    *** Making a personalized recommendation... ***

    You must be logged in to get recommendations

    1. Create an account
    2. Log in
    3. Rate books
    4. Display a generic list of the 3 most popular books
    5. Get a generic book reccommendation
    6. Get a personalized book recommendation

    Enter an option(1-6), or Q to quit: 1

    *** Creating an account... ***

    Create your username: Mike
    Mike is already in use, please try a different username
    Create your username: KittyP
    Account was created successfully. Please log in to continue

    1. Create an account
    2. Log in
    3. Rate books
    4. Display a generic list of the 3 most popular books
    5. Get a generic book reccommendation
    6. Get a personalized book recommendation

    Enter an option(1-6), or Q to quit: 2

    *** Logging in... ***

    Type in your username: KittyP
    Login succeeded. 
    Welcome back KittyP

    1. Create an account
    2. Log in
    3. Rate books
    4. Display a generic list of the 3 most popular books
    5. Get a generic book reccommendation
    6. Get a personalized book recommendation

    Enter an option(1-6), or Q to quit: 3

    *** Rating books... ***


    The Hitchhiker's Guide To The Galaxy (1989) by Douglas Adams

    You have not rated this book
    To rate this book please enter a score (1-5)
    To leave the score unchanged press 'Enter'
    To quit press 'q'
    3

    Watership Down (1978) by Richard Adams

    You have not rated this book
    To rate this book please enter a score (1-5)
    To leave the score unchanged press 'Enter'
    To quit press 'q'
    6
    Invalid Input
    You have not rated this book
    To rate this book please enter a score (1-5)
    To leave the score unchanged press 'Enter'
    To quit press 'q'
    3

    The Five People You Meet in Heaven (1988) by Mitch Albom

    You have not rated this book
    To rate this book please enter a score (1-5)
    To leave the score unchanged press 'Enter'
    To quit press 'q'
    1

    Speak (2009) by Laurie Halse Anderson

    You have not rated this book
    To rate this book please enter a score (1-5)
    To leave the score unchanged press 'Enter'
    To quit press 'q'

    I do not want to rate this book

    I Know Why the Caged Bird Sings (1989) by Maya Angelou

    You have not rated this book
    To rate this book please enter a score (1-5)
    To leave the score unchanged press 'Enter'
    To quit press 'q'
    q

    1. Create an account
    2. Log in
    3. Rate books
    4. Display a generic list of the 3 most popular books
    5. Get a generic book reccommendation
    6. Get a personalized book recommendation

    Enter an option(1-6), or Q to quit: 3

    *** Rating books... ***


    The Hitchhiker's Guide To The Galaxy (1989) by Douglas Adams

    You have rated this book 3
    To rate this book please enter a score (1-5)
    To delete your score press 'd'
    To leave the score unchanged press 'Enter'
    To quit press 'q'
    d

    Watership Down (1978) by Richard Adams

    You have rated this book 3
    To rate this book please enter a score (1-5)
    To delete your score press 'd'
    To leave the score unchanged press 'Enter'
    To quit press 'q'

    I want to keep my current rating

    The Five People You Meet in Heaven (1988) by Mitch Albom

    You have rated this book 1
    To rate this book please enter a score (1-5)
    To delete your score press 'd'
    To leave the score unchanged press 'Enter'
    To quit press 'q'

    I want to keep my current rating

    Speak (2009) by Laurie Halse Anderson

    You have not rated this book
    To rate this book please enter a score (1-5)
    To leave the score unchanged press 'Enter'
    To quit press 'q'
    d
    Invalid Input
    You have not rated this book
    To rate this book please enter a score (1-5)
    To leave the score unchanged press 'Enter'
    To quit press 'q'
    q

    1. Create an account
    2. Log in
    3. Rate books
    4. Display a generic list of the 3 most popular books
    5. Get a generic book reccommendation
    6. Get a personalized book recommendation

    Enter an option(1-6), or Q to quit: 4

    *** Displaying top 3... ***

    ---------------------------------------------------
    The three highest rated books in our database are. 
    ---------------------------------------------------
    Sabriel (2009) by Garth Nix
    Average Rating:  4.0

    The Bourne Series (2004) by Robert Ludlum
    Average Rating:  3.86

    The Hitchhiker's Guide To The Galaxy (1989) by Douglas Adams
    Average Rating:  3.83


    1. Create an account
    2. Log in
    3. Rate books
    4. Display a generic list of the 3 most popular books
    5. Get a generic book reccommendation
    6. Get a personalized book recommendation

    Enter an option(1-6), or Q to quit: 5

    *** Making a generic recommendation... ***

    You might like: 
       Sabriel (2009) by Garth Nix
       Average Rating:  4.0

    1. Create an account
    2. Log in
    3. Rate books
    4. Display a generic list of the 3 most popular books
    5. Get a generic book reccommendation
    6. Get a personalized book recommendation

    Enter an option(1-6), or Q to quit: 6

    *** Making a personalized recommendation... ***

    You might like:
      The Great Gatsby (1925) by F. Scott Fitzgerald
      Rated 5 by Tony

    1. Create an account
    2. Log in
    3. Rate books
    4. Display a generic list of the 3 most popular books
    5. Get a generic book reccommendation
    6. Get a personalized book recommendation

    Enter an option(1-6), or Q to quit: q

    *** Logging out... ***


    -----------------------
     HAVE A NICE DAY. BYE 
    -----------------------

```python

```

```python

```